# Python Basics

## Datentypen und Datencontainer


```python
%pip install -q otter-grader

import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Aufgabe 1:** Erstellen Sie zwei Variablen `income_person1` und `income_person2` und weisen Sie ihnen die Werte 100 und 80 zu.

```python
income_person1 = 100
income_person2 = ...
```

```python
grader.check("q1")
```

**Aufgabe 2:** Im Folgenden werden zwei Variablen `a` und `b` definiert. Berechnen Sie den Mittelwert der beiden Variablen und speichern Sie das Ergebnis in der Variable `mean_ab`. Benutzen Sie dabei nur Addition (`+`) und Division (`/`).

```python
a = 123
b = 55

mean_ab = ...
mean_ab
```

```python
grader.check("q2")
```

**Aufgabe 3:** Wie viel ist 3.12 hoch 8? Berechnen Sie das Ergebnis und speichern Sie es in der Variable `result`.

```python
result = ...
result
```

```python
grader.check("q3")
```

**Aufgabe 4:** Speichern Sie die Zeichenkette "Jane Doe Junior" in der Variable `person_name`.

```python
person_name = ...
person_name
```

```python
grader.check("q4")
```

**Aufgabe 5:** Wie viele Zeichen hat die Zeichenkette in der Variable `person_name`? Speichern Sie das Ergebnis in der Variable `name_length`. Benutzen Sie dazu die Funktion `len()`. Als Beispiel: `len("Hallo")` gibt 5 zurück.

```python
name_length = ...
```

```python
grader.check("q5")
```

** Aufgabe 6 **: Erstellen Sie eine Liste mit den Zahlen 112.3, 302.7, 400.0, und 100.23. Speichern Sie diese
in einer Variable `gas_bills_usd`.

**Hinweis**: Eine Liste können Sie mit eckigen Klammern `[]` erstellen, y.B. `my_list = [1, 2, 3]`.

Erstellen Sie dann eine leere Liste `gas_bills_eur` und fügen Sie die Werte aus 
`gas_bills_usd` in Euro umgerechnet hinzu. Der Umrechnungskurs sei 0.85 USD für 1 EUR.

**Hinweis**: Sie können Werte zu einer Liste mit der Methode `append` hinzufügen, y.B. `my_list.append(4)`.

```python
empty_list = []
for element in another_list:
    empty_list.append(element)
```

```python
gas_bills_usd = []

gas_bills_eur = []
for bill in gas_bills_usd:
    gas_bills_eur.append(...)
gas_bills_eur
```

```python
grader.check("q6")
```

** Aufgabe 7 **: In dieser Aufgabe bearbeiten Sie eine Liste von Nutzungszeiten von Cloud-Computing-Ressourcen. Die Liste `usage` enthält die Nutzungszeiten in Minuten. Wir möchten die Nutzungszeiten, die über ein Limit von 50 Stunden hinausgehen in einer neuen Liste `usage_limit_exceeded` speichern. Benutzen Sie dazu eine List-Komprehension. In dem Beispiel unten wird die Liste `list2` aus den Elementen von `list1` erstellt, die größer als 2 sind.

```python
list1 = [1, 2, 5, 2]
list2 = [element for element in list1 if element > 2]
```

```python
usage = [4839, 2394, 5839, 2394, 802, 904]
usage_limit_exceeded = [minutes for ... in usage if ...]

```

```python
grader.check("q7")
```

** Aufgabe 8 **: In dieser Aufgabe bearbeiten wir eine Liste von Überlebenszeiten von elektrischen Glühbirnen. Die Liste `bulb_survival` enthält die Überlebenszeiten in 1000 Stunden. Wir haben festgestellt, daß die gemessenen Zeiten die tatsächlichen Überlebenszeiten um 10% überschätzen. Erstellen Sie eine neue Liste `bulb_survival_corrected` mit den korrigierten Überlebenszeiten. Benutzen Sie dazu eine List-Komprehension.

```python

```python
bulb_survival = [25.2, 30.1, 35.2, 40.3, 45.2]
bulb_survival_corrected = [survival * ... for survival in ...]
```

```python
grader.check("q8")
```

** Aufgabe 9 **: In dieser Aufgabe bearbeiten wir wieder eine Liste von Überlebenszeiten von elektrischen Glühbirnen. Die Liste `bulb_survival1` enthält die Überlebenszeiten in 1000 Stunden. Wir haben festgestellt, dass die tatsächliche Messung
für die zweite Glühbirne war fehlerhaft war. Die tatsächliche Überlebenszeit war 1200 Stunden. Ändern Sie die Liste `bulb_survival1` so, dass die zweite Überlebenszeit korrigiert wird.

**Hinweis**: Sie können den Wert eines Elements in einer Liste mit `my_list[index]` ändern.

In der Liste list1 unten wird das erste Element auf 100 gesetzt:

```python
list1 = [1, 2, 5, 2]
list1[0] = 100
```

```python

bulb_survival1 = [25.2, 30.1, 35.2, 40.3, 45.2]
bulb_survival1[...] = ...
```

```python
grader.check("q9")
```

**Aufgabe 10**: In dieser Aufgabe haben aus unserer Kundendatenbank die Daten einer Person heruntergeladen und in einer Dictionary `best_customer` mit Schlüsseln
 `name`, `age`, `city`, und `total_purchases` gespeichert. Wir möchten eine Nachricht an den Kunden senden, in der wir den Namen des Kunden und den Umsatz erwähnen. Erstellen Sie eine Zeichenkette `msg` mit der Nachricht. Benutzen Sie dazu die Werte aus dem Dictionary `best_customer`.

**Hinweis**: Sie können auf die Werte eines Dictionarys mit `my_dict[key]` zugreifen.

```python
my_dict = {"name": "John", "age": 30}
print(my_dict["name"]) # gibt "John" aus
```

**Hinweis**: In der formattierten Zeichenkette können Sie Werte mit `{}` einfügen. 

```python
name = "John"
age = 30
msg = f"Name: {name}, Age: {age}"
print(msg) # gibt "Name: John, Age: 30" aus
```

**Hinweis**: In der formattierten Zeichenkette können Sie auch auf die Werte eines Dictionarys zugreifen.

```python
my_dict = {"name": "John", "age": 30}
msg = f"Name: {my_dict['name']}, Age: {my_dict['age']}"
print(msg) # gibt "Name: John, Age: 30" aus
```

Achten Sie darauf, dass die Anführungszeichen in der formattierten Zeichenkette korrekt gesetzt sind: benutzen Sie `'`, um die Schlüssel des Dictionarys zu umschließen, und `"` für die formattierte Zeichenkette.

```python
best_customer = {
    "name": "Phillip Jones",
    "age": 35,
    "total_purchases": 30010.5
}
msg = f"Dear {...}, your total purchases amount to {...} EUR."
```

```python
grader.check("q10")
```

**Aufgabe 11**: Erstellen Sie ein Dictionary `customer` mit den Schlüsseln `first_name`, `last_name`, `city`, `session_durations` und den Werten "Jane Doe", 30, "New York", [2.4, 8.1, 2.2]. Under `session_durations` haben wir die Dauer der letzten drei Sitzungen (z.B. Softwarenutzung) in Stunden gespeichert.

Setzen Sie die Dauer der letzten Sitzung auf 0 Stunden.

**Hinweis**: Sie können den Wert eines Elements in einer Liste mit `my_list[index]` ändern.

In der Liste list1 unten wird das erste Element auf 100 gesetzt:

```python
list1 = [1, 2, 5, 2]
# Wird das erste Element auf 100 setzen
list1[0] = 100
```

Falls die Liste unter einem Schlüssel in einem Dictionary gespeichert ist, können Sie auf die Liste mit `my_dict[key]` zugreifen.

```python
my_dict = {"name": "John", "age": 30, "some_list": [1, 2, 8]}
# Wird das erste Element der Liste auf 0 setzen
my_dict["session_durations"][0] = 0
```



```python
customer = {
    "first_name": ...,
    "last_name": ...,
    "city": ...,
    ...
}
customer[...][...] = 0
```

```python
grader.check("q11")
```

**Aufgabe 12**: In dieser Aufgabe werden wir eine Liste aus Dictionarys bearbeiten. Wir können uns vorstellen, dass diese Liste aus einer Datenbankabfrage stammt. Jedes Dictionary in der Liste enthält die Daten eines Kunden. Wir möchten nun eine Schleife verwenden, um eine Liste `customer_names_new_york` mit den Namen der Kunden zu erstellen, die in New York leben.

```python

```

```python
customers = [
    {"name": "John", "age": 30, "city": "New York"},
    {"name": "Jane", "age": 25, "city": "Los Angeles"},
    {"name": "Tom", "age": 35, "city": "Chicago"},
    {"name": "Tim", "age": 40, "city": "Houston"},
    {"name": "Angelina", "age": 45, "city": "Phoenix"}
    {"name": "Tina", "age": 45, "city": "New York"}
]

customers_names_new_york = []

for customer in customers:
    if ...:
        customers_names_new_york.append(...)
```

```python
grader.check("q12")
```

**Aufgabe 13**: In dieser Aufgabe werden wir eine Liste aus Listen bearbeiten. Der Sinn der Aufgabe ist es, Ihnen zu zeigen, wie Variablen in Python funktionieren. Die Liste besteht aus drei anderen Listen. Wir möchten nun die erste Liste einer Variablen `first_list` zuweisen. Danach ändern wir das erste Element von `first_list` auf 100. Was passiert mit der ursprünglichen Liste `list_of_lists`?

```python
list_of_lists = [[1, 2, 3], [4, 5, 6], [7, 8, 9]]
first_list = list_of_lists[0]
first_list[0] = 100

list_of_lists
```

**Aufgabe 13a**: Im folgenden Code wird eine Liste `accounts` mit den Werten 33 und 80 definiert. Weisen Sie diese zwei Werte den Variablen `account1` und `account2` zu, indem Sie die Werte aus der Liste `accounts` entpacken.

```python
ids = [1, 2, 3]
id1, id2, id3 = ids
```

```python
accounts = [33, 80]
account1, ... = accounts
```

```python
grader.check("q13a")
```

### Tupel

Bisher haben wir Listen verwendet, um eine Sammlung von Elementen zu speichern. Eine Liste ist veränderbar, d.h. wir können die Elemente der Liste verändern. Ein Tupel ist eine Sammlung von Elementen, die nicht veränderbar ist. Ein Tupel wird mit runden Klammern `()` erstellt.

```python
my_list = [1, 4]
my_tuple = (1, 4)

my_list[0] = 100
# my_tuple[0] = 100
```

## Funktionen

In allen Programmiersprachen sind Funktionen ein wichtiges Konzept. Funktionen sind ein Weg, um Code zu organisieren und zu strukturieren.

Als Beispiel betrachten wir die Funktion `add_numbers`. Diese Funktion nimmt zwei Argumente `a` und `b` und gibt die Summe der beiden Argumente zurück.

```python
def add_numbers(a, b):
    return a + b
```

Um eine Funktion zu definieren, benutzen wir das Schlüsselwort `def`, gefolgt von dem *Namen* der Funktion und den Argumenten in Klammern. Der Funktionskörper ist eingerückt. Die Funktion gibt das Ergebnis mit dem Schlüsselwort `return` zurück.

Um die Funktion aufzurufen, benutzen wir den Namen der Funktion gefolgt von den Argumenten in Klammern.

```python
result = add_numbers(3, 5)
print(result) # gibt 8 aus
```

** Aufgabe 14 **: Erstellen Sie eine Funktion `calculate_mean` die den Mittelwert von zwei Zahlen berechnet. Die Funktion soll zwei Argumente `a` und `b` haben und den Mittelwert der beiden Argumente zurückgeben.

Der Mittelwert von zwei Zahlen ist die Summe der beiden Zahlen geteilt durch 2.
$$
\text{mean} = \frac{a + b}{2}
$$


```python
def calculate_mean(a, ...):
    return ...
```

```python
grader.check("q14")
```

**Aufgabe 15**: In der vorherigen Aufgabe haben wir die Funktion `calculate_mean` erstellt. In dieser Aufgabe möchten wir die Funktion erweitern, so dass sie den Mittelwert von beliebig vielen Zahlen berechnen kann. Die Funktion soll eine Liste von Zahlen als Argument nehmen und den Mittelwert der Zahlen zurückgeben. Benutzen Sie dazu die Funktion `sum` und die Funktion `len`.

```python
def calculate_mean_list(...):
    return sum(...) / ...
```

```python
grader.check("q15")
```

**Aufgabe 16**: In dieser Aufgabe möchten wir eine Funktion `average_word_length` definieren. Die Funktion soll eine Zeichenkette als Argument nehmen und die durchschnittliche Länge der Wörter in der Zeichenkette zurückgeben. Wir werden annehmen, dass Wörter durch Leerzeichen getrennt sind. Benutzen Sie dazu die Methode `split` und die Funktion `len`.

**Hinweis**: Die Methode `split` gibt eine Liste aus Zeichenketten, die durch Trennzeichen getrennt sind, zurück.

```python
my_string = "Hello World"
words = my_string.split(" ")
print(words) # gibt ["Hello", "World"] aus
```

```python
my_string = "Cheese, Bread, Milk        and Butter"
words = my_string.split(", ")
print(words) # gibt ["Cheese", "Bread", "Milk        and Butter"] aus
```

```python
def average_word_length(sentence):
    words = ...
    word_lengths = [len(...) for ... in ...]
    return ...
```

```python
average_word_length("Hello World") == 5
```

```python
average_word_length("Hi everybody") == (9 + 2) / 2
```

Tests for q1 (points: None)
```
>>> income_person1 == 100
True
```
```
>>> income_person2 == 80
True
```


Tests for q10 (points: None)
```
>>> msg == 'Dear Phillip Jones, your total purchases amount to 30010.5 EUR.'
True
```


Tests for q11 (points: None)
```
>>> customer['first_name'] == 'John'
True
```
```
>>> customer['last_name'] == 'Davis'
True
```
```
>>> customer['city'] == 'New York'
True
```
```
>>> len(customer['session_durations']) == 3
True
```
```
>>> customer['session_durations'][-1] == 0
True
```


Tests for q12 (points: None)
```
>>> len(customers_names_new_york) == 2
True
```
```
>>> customers_names_new_york[0] == 'John'
True
```
```
>>> customers_names_new_york[1] == 'Tina'
True
```


Tests for q13a (points: None)
```
>>> account1 == 33
True
```
```
>>> account2 == 80
True
```


Tests for q14 (points: None)
```
>>> calculate_mean(0, 0) == 0
True
```
```
>>> calculate_mean(1, 1) == 1
True
```


Tests for q15 (points: None)
```
>>> calculate_mean_list([0, 0]) == 0
True
```
```
>>> calculate_mean_list([1, 1]) == 1
True
```


Tests for q2 (points: None)
```
>>> mean_ab == 89
True
```


Tests for q3 (points: None)
```
>>> np.isclose(result, 3.12 ** 8)
True
```


Tests for q4 (points: None)
```
>>> person_name == 'Jane Doe Junior'
True
```


Tests for q5 (points: None)
```
>>> name_length == len(person_name_2)
True
```


Tests for q6 (points: None)
```
>>> len(gas_bills_usd) == 4
True
```
```
>>> len(gas_bills_eur) == len(gas_bills_usd)
True
```
```
>>> assert np.isclose(gas_bills_eur[0], gas_bills_usd[0] * 0.85)
```


Tests for q7 (points: None)
```
>>> len(usage_limit_exceeded) == 2
True
```
```
>>> usage_limit_exceeded[0] == 4839
True
```


Tests for q8 (points: None)
```
>>> len(bulb_survival_corrected) == len(bulb_survival)
True
```
```
>>> bulb_survival_corrected[0] == bulb_survival[0] * 0.9
True
```


Tests for q9 (points: None)
```
>>> len(bulb_survival1) == 5
True
```
```
>>> bulb_survival1[1] == 1.2
False
```
